<!-- DEBUT problem Les normes et leurs sigles -->
<problem display_name="Les normes et leurs sigles" >
    <legend>
        <p>
 <strong>
  Classez ces modes de connexion du plus lent au plus rapide.
 </strong>
</p>

<p>
 3G,4G,H+,Edge
</p>
    </legend>
    
    <script type="loncapa/python">
import re
def checkAnswerEssay(expect, ans):
    response = re.search('', ans)
    if response:
        return 1
    else:
        return 0
</script>
<span id="36e097dd-2686-4b6a-a59f-e6f6356059d0"></span>
<script type="text/javascript">
    /* The object here is to replace the single line input with a textarea */
   (function() {
    var elem = $("#36e097dd-2686-4b6a-a59f-e6f6356059d0")
        .closest("div.problem")
        .find(":text");
    /* There's CSS in the LMS that controls the height, so we have to override here */
    var textarea = $('<textarea style="height:150px" rows="20" cols="70" />');
    console.log(elem);
    console.log(textarea);
    //This is just a way to do an iterator in JS
    for (attrib in {'id':null, 'name':null}) {
        textarea.attr(attrib, elem.attr(attrib));
    }
    /* copy over the submitted value */
    textarea.val(elem.val())
    elem.replaceWith(textarea);
    
   })();
</script>
    <customresponse cfn="checkAnswerEssay">
    	<textline size="40" correct_answer="" label="Problem Text"/>
    </customresponse>
    
    
    <solution>
        <div class="detailed-solution"><h1>
 Les normes et leurs sigles
</h1>

<ul>
 <li>
  Les modes de connexion du plus lent au plus rapide.
  <ul>
   <li>
    E (Edge) aussi appelé 2G, lent. Ce mode de connexion permet à peine de lire ses mails. Il ne permet pas une navigation fluide sur le Web.
   </li>
   <li>
    3G (3ème génération) permet de faire des recherches et de surfer sans trop attendre.
   </li>
   <li>
    H+, est une amélioration de la 3G. il est plus rapide que le wifi si les connexions sont optimales. Et l'accès à la musique en ligne où aux vidéos peut être envisagé.
   </li>
   <li>
    4G, plus rapide que le wifi si les connexions sont optimales. À condition bien sûr que cette connexion soit de bonne qualité ("plusieurs petites briques"), l'accès à internet est alors très fluide, et les jeux en ligne, les vidéos en streaming ou le téléchargement de gros fichiers devient possible.
   </li>
  </ul>
 </li>
</ul>

<p>
 Notez bien que pour pouvoir bénéficier d'une connexion 4G, il faut :
</p>

<ul>
 <li>
  que cette connexion soit disponible là où vous vous trouvez,
 </li>
 <li>
  que votre smartphone soit équipé d'une antenne 4G, c'est loin d'être le cas sur tous les modèles y compris sur des appareils récents.
 </li>
</ul></div>
    </solution>
    
</problem>
<!-- FIN problem Les normes et leurs sigles -->
