<problem display_name="La charte informatique de l'université" max_attempts=""><legend><p><strong>La charte informatique de l'université</strong><br />
Pour utiliser les <em>services numériques</em> de l'université, vous avez signé une <strong>charte informatique</strong> lors de votre inscription :<br />
Quel paragraphe de la charte <a href="http://www.univ-lille3.fr/portail/charte-informatique/">http://www.univ-lille3.fr/portail/charte-informatique/</a> met en lumière la question du suivi des informations des utilisateurs.<br />
Si vous vous connectez avec la <strong>4G</strong> depuis l'université, devez-vous respecter cette charte ? Et avec votre smartphone et une connexion wifi ?</p> </legend><script type="loncapa/python">
import re
def checkAnswerEssay(expect, ans):
    response = re.search('', ans)
    if response:
        return 1
    else:
        return 0
            </script><span id="45e61def-0304-49b7-9937-132aabd015d5"></span><script type="text/javascript">
    /* The object here is to replace the single line input with a textarea */
    (function() {
    var elem = $("#45e61def-0304-49b7-9937-132aabd015d5")
        .closest("div.problem")
        .find(":text");
    /* There's CSS in the LMS that controls the height, so we have to override here */
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>');
    console.log(elem);
    console.log(textarea);
    //This is just a way to do an iterator in JS
    for (attrib in {'id':null, 'name':null}) {
        textarea.attr(attrib, elem.attr(attrib));
    }
    /* copy over the submitted value */
    textarea.val(elem.val())
    elem.replaceWith(textarea);

    })();
            </script><customresponse cfn="checkAnswerEssay"><textline size="40" correct_answer="" label="Problem Text"/></customresponse><solution><div class="detailed-solution"><p>1/ Quel paragraphe de la charte <a href="http://www.univ-lille3.fr/portail/charte-informatique/">http://www.univ-lille3.fr/portail/charte-informatique/</a> met en lumière la question du filtrage et du suivi ?<br />
<em>Les engagements de l'Université</em><br />
<code>Réglementairement, l'Université recueille et conserve les informations sur les accès à ses services numériques. L'Université peut être obligée de fournir ces informations lors d'une enquête judiciaire.</code><br />
2/ Si vous vous connectez avec la <strong>4G</strong> depuis l'université, devez-vous respecter cette charte ? Et avec votre smartphone et une connexion wifi ?<br />
OUI et OUI, les services numériques en particulier ceux accessibles depuis <strong>l'ENT</strong> sont soumis à cette charte <strong>quel que soit le mode de connexion</strong>. En revanche, si vous vous connectez en 4G à d'autres sites que ceux de l'université tout en étant physiquement dans les batiments de l'établissement, aucune information ne sera collectée par l'université.</p> </div></solution></problem>