<problem display_name="Noms de domaine" max_attempts=""><legend><p><strong>Comment s'appelle l'organisme qui gère les noms de domaine et les adresses IP ?</strong> Précisez votre réponse pour la France et au niveau mondial.</p> </legend><script type="loncapa/python">
import re
def checkAnswerEssay(expect, ans):
    response = re.search('', ans)
    if response:
        return 1
    else:
        return 0
            </script><span id="514dfad2-8396-451f-a1ab-701f73d3a194"></span><script type="text/javascript">
    /* The object here is to replace the single line input with a textarea */
    (function() {
    var elem = $("#514dfad2-8396-451f-a1ab-701f73d3a194")
        .closest("div.problem")
        .find(":text");
    /* There's CSS in the LMS that controls the height, so we have to override here */
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>');
    console.log(elem);
    console.log(textarea);
    //This is just a way to do an iterator in JS
    for (attrib in {'id':null, 'name':null}) {
        textarea.attr(attrib, elem.attr(attrib));
    }
    /* copy over the submitted value */
    textarea.val(elem.val())
    elem.replaceWith(textarea);

    })();
            </script><customresponse cfn="checkAnswerEssay"><textline size="40" correct_answer="" label="Problem Text"/></customresponse><solution><div class="detailed-solution"><p>ICANN pour le niveau mondial et l'AFNIC pour la France.</p> </div></solution></problem>