<problem display_name="Le numérique concerne tout le monde" max_attempts=""><legend><p><strong>Quels étudiants sont concernés par le numérique ?</strong></p>
<p>Le numérique concerne évidemment les étudiants en informatique et plus généralement les étudiants des filières scientifiques.  Mais vous qui êtes inscrits dans une université de sciences humaines et sociales, êtes-vous concernés ?</p>
<p>Choisissez au moins 3 des domaines suivants et faites des recherches pour voir en quoi ils sont impactés par le numérique : les médias, la santé, l'histoire, la sociologie, la linguistique, les arts, la culture, l'enseignement, l'archéologie.</p>
<p>Faites une synthèse en quelques lignes de vos recherches en précisant les domaines auxquels vous vous êtes intéressés. Indiquez les liens des sites sur lesquels vous avez trouvé ces informations. La liste est non exhaustive et vous pouvez vous intéresser à d'autres domaines.</p> </legend><script type="loncapa/python">
import re
def checkAnswerEssay(expect, ans):
    response = re.search('', ans)
    if response:
        return 1
    else:
        return 0
            </script><span id="6655a16a-9c09-4ad4-8fd3-93a7e6ce5f81"></span><script type="text/javascript">
    /* The object here is to replace the single line input with a textarea */
    (function() {
    var elem = $("#6655a16a-9c09-4ad4-8fd3-93a7e6ce5f81")
        .closest("div.problem")
        .find(":text");
    /* There's CSS in the LMS that controls the height, so we have to override here */
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>');
    console.log(elem);
    console.log(textarea);
    //This is just a way to do an iterator in JS
    for (attrib in {'id':null, 'name':null}) {
        textarea.attr(attrib, elem.attr(attrib));
    }
    /* copy over the submitted value */
    textarea.val(elem.val())
    elem.replaceWith(textarea);

    })();
            </script><customresponse cfn="checkAnswerEssay"><textline size="40" correct_answer="" label="Problem Text"/></customresponse><solution><div class="detailed-solution"><h1>Le numérique concerne tout le monde</h1>
<p>Ces recherches ont dû vous convaincre, si c'était nécessaire, que le numérique <strong>n'est pas réservé</strong> aux informaticiens, il concerne tout le monde, toutes les disciplines.<br />
S'agissant plus particulièrement des <strong>sciences humaines</strong>, la prise en compte du numérique a fait évoluer les champs disciplinaires pour faire apparaître ce qu'on appelle les <strong>humanités numériques</strong> ( <em>digital humanities</em> en anglais).</p>
<p>Voici quelques exemples que nous vous proposons, n'hésitez pas à proposer d'autres exemples dans le forum de discussion :</p>
<ul>
<li>Dans les <strong>médias</strong> : nouveau sous-métier de journalisme : les <strong>data-journalistes</strong><ul>
<li><a href="http://www.lemonde.fr/data-visualisation/">data-visualisation</a></li>
<li><a href="http://fr.wikipedia.org/wiki/Journalisme_de_données">journalisme de données</a></li>
</ul>
</li>
<li>Dans la <strong>santé</strong> : (imagerie, dossier numérique du patient, ...)<ul>
<li><a href="https://interstices.info/jcms/c_21525/simulation-de-loperation-de-la-cataracte">simulation</a></li>
</ul>
</li>
<li>En <strong>histoire, sociologie, linguistique</strong> : <em>fouille de données</em><ul>
<li><a href="http://www.youtube.com/watch?feature=player_embedded&amp;v=tp4y-_VoXdA">fouille de données</a></li>
</ul>
</li>
<li>En <strong>art et culture</strong> :<ul>
<li><a href="http://www.lefresnoy.net/fr/Le-Fresnoy/presentation">Le Fresnoy</a></li>
</ul>
</li>
<li>Dans l'<strong>enseignement</strong> : (outils numérique d'accompagnement scolaire, MOOC,...):<ul>
<li><a href="https://www.france-universite-numerique-mooc.fr/cours/">FUN</a></li>
</ul>
</li>
<li>En fouille archéologique :  une réalisation prestigieuse réalisée à Lille3 :<ul>
<li><a href="http://bsa.biblio.univ-lille3.fr/blog/2013/09/exposition-le-vase-qui-parle-au-palais-des-beaux-arts-de-lille/">vase qui parle</a></li>
</ul>
</li>
</ul> </div></solution></problem>