<problem display_name="Utiliser les serveurs de noms" max_attempts=""><legend><p><strong>Utiliser les serveurs de noms</strong></p>
<p>Il existe des <em>outils</em> dans votre ordinateur pour interroger les serveurs de noms et faire donc la <em>conversion</em> entre <code>adresses IP</code> et <code>noms</code>. Ils sont intégrés par exemple dans le navigateur web et la plupart des outils qui utilisent internet. Même s'il est possible de les utiliser directement, ils sont assez techniques.<br />
Nous vous proposons d'utiliser plutôt un <strong>service</strong> disponible sur internet à travers le navigateur : <a href="http://www.monip.fr">http://www.monip.fr</a><br />
Selon ce site :<br />
  - quelle est votre adresse IP ?<br />
  - où vous trouvez-vous ?</p>
<p>D'après ce site, où se trouvent les machines suivantes :</p>
<ul>
<li><a href="www.wikipedia.org">www.wikipedia.org</a></li>
<li><a href="www.facebook.com">www.facebook.com</a></li>
<li><a href="www.google.com">www.google.com</a></li>
</ul>
<p>Quelle est l'adresse IP de la machine qui porte le nom <a href="www.grappa.univ-lille3.fr">www.grappa.univ-lille3.fr</a> ?</p>
<p><strong>Copiez</strong> l'adresse IP que vous venez de  trouver dans la <strong>barre d'adresse de votre navigateur.</strong> Dans un autre onglet, <strong>copiez</strong> le <strong>nom</strong> : www.grappa.univ-lille3.fr dans la barre d'adresse.</p>
<ul>
<li>Que constatez vous ?</li>
</ul>
<p>Une fois validé votre rendu, vous pourrez consulter la correction pour vous auto-évaluer, ...</p> </legend><script type="loncapa/python">
import re
def checkAnswerEssay(expect, ans):
    response = re.search('', ans)
    if response:
        return 1
    else:
        return 0
            </script><span id="921f55df-c66e-499d-b75c-14b7cb0a3687"></span><script type="text/javascript">
    /* The object here is to replace the single line input with a textarea */
    (function() {
    var elem = $("#921f55df-c66e-499d-b75c-14b7cb0a3687")
        .closest("div.problem")
        .find(":text");
    /* There's CSS in the LMS that controls the height, so we have to override here */
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>');
    console.log(elem);
    console.log(textarea);
    //This is just a way to do an iterator in JS
    for (attrib in {'id':null, 'name':null}) {
        textarea.attr(attrib, elem.attr(attrib));
    }
    /* copy over the submitted value */
    textarea.val(elem.val())
    elem.replaceWith(textarea);

    })();
            </script><customresponse cfn="checkAnswerEssay"><textline size="40" correct_answer="" label="Problem Text"/></customresponse><solution><div class="detailed-solution"><h2>Utiliser les serveurs de noms</h2>
<h3>Votre propre IP</h3>
<p>Grâce à <a href="http://www.monip.fr">http://www.monip.fr</a> vous avez trouvé votre adresse IP, elle est évidemment différente pour chacun d'entre nous.<br />
Si vous recommencez depuis une autre machine ou avec la même machine connectée différemment (wifi domicile ou wifi eduroam, connexion filaire à l'université ou eduroam université, connexion wifi ou connexion 3G/4G) vous obtiendrez à chaque fois un résultat <strong>différent</strong>.</p>
<p>Pour ce qui est de la <strong>localisation</strong>, vous avez sûrement obtenu une adresse géographiquement proche de la votre, celle-ci est en fait l'<em>adresse du  routeur du fournisseur d'accès</em> du réseau auquel vous êtes connecté.<br />
Vous pouvez également constater que vos informations personnelles telles que :</p>
<ul>
<li>le navigateur que vous utilisez,</li>
<li>la langue utilisée,</li>
<li>la taille de votre écran, ...</li>
</ul>
<p>ont été <strong>repérés</strong> par le site www.monip.fr et vous sont renvoyés dans la liste des informations.</p>
<h3>Wikipedia, facebook et google</h3>
<p>Vous avez dû trouver que la localisation de wikipedia est aux <em>Pays-Bas</em>, près <em>d'Amsterdam</em>. C'est en effet là que se trouve la copie la plus proche.<br />
D'autres répliques de la célèbre encyclopédie existent également en <em>Floride</em> et en <em>Corée du  Sud</em>.<br />
<code>Facebook</code> est hébergé à <em>Kensington</em>, dans <em>l'Ohio</em> aux <em>États-Unis</em>, c'est là que les <em>données</em> de tous les utilisateurs sont stockées.<br />
Alors que <code>Google</code> se trouve à <em>Mountain View</em> en <em>Californie</em>.<br />
Toutes les <strong>lois</strong> concernant la sécurisation de nos données et les respect de <strong>notre vie privée</strong> pour ces 2 sites sont donc des lois <strong>américaines</strong>.</p>
<h3>www.monip.fr</h3>
<p>et oui, ce site est une version traduite en français d'un site Allemand.</p>
<h3>www.grappa.univ-lille3.fr</h3>
<p>l'adresse <strong>IP</strong> de l'équipe de recherche Grappa est : <code>194.254.132.190</code><br />
Si vous copiez cette adresse dans dans la barre d'adresse de votre navigateur, vous tombez exactement sur la <strong>même page</strong> que si vous allez sur <a href="www.grappa.univ-lille3.fr">www.grappa.univ-lille3.fr</a>.<br />
Cela ne doit pas vous surprendre si vous avez compris l'objet de cette séquence.</p> </div></solution></problem>