<problem display_name="Le numérique au quotidien" max_attempts=""><legend><p>Les microprocesseurs, les ordinateurs ont envahi notre quotidien. Pour chacun des domaines suivants, cherchez des exemples où le numérique a permis des évolutions notables :</p>
<ul>
<li>Domotique</li>
<li>Transports</li>
<li>Vêtements</li>
<li>Médical / paramédical</li>
</ul>
<p>Après avoir effectué vos recherches, copier dans la fenêtre de rendu 1 lien pour au moins 3 des 4 thèmes proposés (un lien par thème).</p> </legend><script type="loncapa/python">
import re
def checkAnswerEssay(expect, ans):
    response = re.search('', ans)
    if response:
        return 1
    else:
        return 0
            </script><span id="ae838c27-732a-4f45-bab4-d9a0cdf17b68"></span><script type="text/javascript">
    /* The object here is to replace the single line input with a textarea */
    (function() {
    var elem = $("#ae838c27-732a-4f45-bab4-d9a0cdf17b68")
        .closest("div.problem")
        .find(":text");
    /* There's CSS in the LMS that controls the height, so we have to override here */
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>');
    console.log(elem);
    console.log(textarea);
    //This is just a way to do an iterator in JS
    for (attrib in {'id':null, 'name':null}) {
        textarea.attr(attrib, elem.attr(attrib));
    }
    /* copy over the submitted value */
    textarea.val(elem.val())
    elem.replaceWith(textarea);

    })();
            </script><customresponse cfn="checkAnswerEssay"><textline size="40" correct_answer="" label="Problem Text"/></customresponse><solution><div class="detailed-solution"><h1>le numérique au quotidien</h1>
<p>Quelques exemples que nous vous proposons au cas où vous n'auriez rien trouvé, ...</p>
<p>La <strong>domotique</strong> est un domaine en pleine expansion qui vise à équiper numériquement notre maison :</p>
<ul>
<li><a href="https://nest.com/fr/">nest</a></li>
<li><a href="http://fr.wikipedia.org/wiki/Domotique">domotique</a></li>
</ul>
<p>Pour les <strong>transports</strong>, les ordinateurs de bord sont depuis longtemps présents dans les voitures, de plus en plus ils sont responsables de notre sécurité :</p>
<ul>
<li><a href="http://fr.wikipedia.org/wiki/electrostabilisateur_programmé">electrostabilisateur</a></li>
<li><a href="http://fr.wikipedia.org/wiki/Ordinateur_de_bord">ordinateur de bord</a></li>
</ul>
<p>Les <strong>chaussures</strong> : gadget ou réelle innovation ? Ce genre d'objet est de plus en plus présents dans nos vies :</p>
<ul>
<li><a href="http://www.linternaute.com/science/technologie/deja-demain/07/chaussure-intelligente/chaussure-intelligente.shtml">chaussures</a><br />
Les <strong>lentilles pour la vue</strong> ?</li>
<li><a href="http://www.zdnet.fr/actualites/google-apres-les-lunettes-connectees-les-lentilles-pour-le-diabete-39797148.htm">lentilles</a></li>
</ul> </div></solution></problem>