<problem display_name="Wifi ou Données mobiles ?" max_attempts="1"><legend><p><strong>Wifi ou Données mobiles ?</strong></p>
<p>Les possesseurs de tablette ou de smartphone peuvent se connecter à internet via le <strong>Wifi</strong> ou le <strong>réseau téléphonique</strong> (<em>données mobiles 3G/4G</em>). Vaut-il mieux se connecter avec l'un ou avec l'autre ?<br />
La réponse dépend de trois paramètres :</p>
<ul>
<li>la <strong>disponibilité</strong> de ces modes de connexions (s'il n'y a pas de réseau wifi, il sera difficile de vous connectez en wifi…)</li>
<li>le <strong>prix</strong> de la connexion</li>
<li>la <strong>rapidité</strong> de la connexion</li>
</ul>
<p>Imaginons plusieurs situations concrètes :</p>
<ol>
<li>je suis chez moi avec un <em>smartphone</em>, le domicile est équipée d'une <em>Box adsl</em> qui fournit un accès <em>Wifi</em>. Ai-je intérêt à me connecter :<br />
        - avec le réseau de données mobiles (3G ou 4G)<br />
        - via le Wifi de la maison</li>
<li>Je suis à <em>l'université</em> et je veux me connecter avec mon <em>smartphone</em>. Ai-je intérêt à me connecter :<br />
        - avec le wifi de l'université (réseau eduroam)<br />
        - avec mon forfait 3G/4G</li>
<li>Je suis en <em>voiture</em> (ce n'est pas moi qui conduis) et je veux consulter mes <em>emails</em> ou faire une <em>recherche</em> sur le net avec mon <em>smartphone</em>, ai-je intérêt à me connecter :<br />
        - en wifi<br />
        - avec le réseau de données mobiles de mon opérateur téléphonique</li>
</ol> </legend><script type="loncapa/python">
import re
def checkAnswerEssay(expect, ans):
    response = re.search('', ans)
    if response:
        return 1
    else:
        return 0
            </script><span id="c0bfd408-beb8-4f86-9013-978f16ca3cce"></span><script type="text/javascript">
    /* The object here is to replace the single line input with a textarea */
    (function() {
    var elem = $("#c0bfd408-beb8-4f86-9013-978f16ca3cce")
        .closest("div.problem")
        .find(":text");
    /* There's CSS in the LMS that controls the height, so we have to override here */
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>');
    console.log(elem);
    console.log(textarea);
    //This is just a way to do an iterator in JS
    for (attrib in {'id':null, 'name':null}) {
        textarea.attr(attrib, elem.attr(attrib));
    }
    /* copy over the submitted value */
    textarea.val(elem.val())
    elem.replaceWith(textarea);

    })();
            </script><customresponse cfn="checkAnswerEssay"><textline size="40" correct_answer="" label="Problem Text"/></customresponse><solution><div class="detailed-solution"><h2>Wifi ou données mobiles ?</h2>
<h3>Wifi ou données mobiles : Comment choisir ?</h3>
<p>1/ Chez moi avec un <strong>smartphone</strong>, j'ai intérêt à me connecter avec le <strong>wifi</strong> de la <strong>Box adsl</strong>.<br />
En effet, la connexion Wifi a déjà été payée dans le cadre de l'abonnement ADSL, la connexion du smartphone ou de la tablette n'engendre donc pas de frais supplémentaires.<br />
Par contre la connexion au réseau 3G/4G est décomptée du forfait qui est souvent limité.</p>
<p>2/ <strong>L'université</strong> a investi dans le réseau wifi <code>eduroam</code>pour proposer ce service aux usagers, il est donc <em>gratuit</em> et n'entraîne aucun frais de connexion.<br />
En revanche comme dans l'exemple précédent, les forfaits <em>données mobiles</em> des abonnements téléphoniques sont souvent limités et/ou chers.<br />
Le réseau Wifi de l'université s'appelle eduroam et tous les usagers peuvent s'y connecter.</p>
<p>3/ En <strong>voiture</strong>, si ce n'est pas moi qui conduit, je peux me connecter à internet en utilisant la connexion <strong>3G/4G</strong>.<br />
Elle permet de rester connecté sur de grandes distances. La couverture en agglomération et dans les zones de forte densité démographique est en général assez bonne pour l'ensemble des opérateurs,<br />
en revanche dans les zones plus reculées, il est parfois difficile de <em>trouver du réseau</em>. Et là tous les opérateurs n'ont pas la même couverture.<br />
La courte portée des antennes Wifi ne permet pas d'utiliser ce mode connexion lors de déplacements importants.</p>
<h3>Le saviez-vous ? - Le relai Wifi</h3>
<p>On trouve dans certaines villes des bus  proposant un accès Wifi.<br />
Si ceux-ci sont équipés d'une antenne de réception 4G, ils peuvent ensuite "redistribuer" la connexion en Wifi, à l'intérieur du véhicule.<br />
Cela est également <strong>possible</strong> avec des <strong>smartphones</strong> récents. Sur le principe décrit ci-dessus, un smartphone peut se connecter à un réseau 3G/4G et ensuite se comporter comme une borne wifi à laquelle peuvent se connecter d'<strong>autres périphériques</strong>.<br />
Cela est très pratique pour se connecter avec un ordinateur là où seules des connections 3G/4G seraient disponibles ... mais attention à la facture !</p> </div></solution></problem>